```python
# Be sure to use Python 3
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

# Core Language

### 1. Make a dictionary (1 point)

Write a function that returns a dictionary whose keys and values are your name, hometown, department, student status (undergrad/grad), and years at TAMU.

*Don't forget:* You will need to delete the line with `raise NotImplementedError`, write your own solution, and then re-run the cell before running the test cell. Each time you change code in a cell, you  will need to re-run that cell before running any other cells that depend on it!

<!-- grade_id: dictionary, points: 0, grade: false, solution: true -->
```python
def intro_dictionary(name, hometown, department, status, years):
    """Returns a dictionary containing key/value pairs.
    
    Parameters
    ----------
    name : string
        Your name
    hometown : string
        Where are you from?
    department : string
        What department are you in?
    status : string
        Are you an undergrad or graduate student?
    years : float
        How many years have you been at TAMU?
        
    Returns
    -------
    dictionary containing the input information
    
    """
    # YOUR CODE HERE
    return {'name': name, 'hometown': hometown,'department':department,'status':status,'years':years}   
intro_dictionary('Xiao', 'Xiangyang', 'Oceanography', 'Graduate', 1)

```

Output:
```
{'department': 'Oceanography',
 'hometown': 'Xiangyang',
 'name': 'Xiao',
 'status': 'Graduate',
 'years': 1}
```

<!-- grade_id: test_dictionary, points: 1, grade: true, solution: false -->
```python
"""(1 point) Test code for the previous function. This cell should NOT give any errors when it is run."""
assert(isinstance(intro_dictionary('Kristen', 'Seattle', 'Oceanography', 'faculty', 4), dict))
assert(intro_dictionary('Kristen', 'Seattle', 'Oceanography', 'faculty', 4)['name'])
assert(intro_dictionary('Kristen', 'Seattle', 'Oceanography', 'faculty', 4)['hometown'])
assert(intro_dictionary('Kristen', 'Seattle', 'Oceanography', 'faculty', 4)['department'])
assert(intro_dictionary('Kristen', 'Seattle', 'Oceanography', 'faculty', 4)['status'])
assert(intro_dictionary('Kristen', 'Seattle', 'Oceanography', 'faculty', 4)['years'])

print("Success!")
```

Output:
```
Success!
```

### 2. Test for vowels (1 point)

Write a function that takes a character (i.e. a string of length 1) and returns True if it is a vowel, False otherwise. For example,
    >>> isVowel("a")
    True
    >>> isVowel("b")
    False

You can compare the input character with a known list of vowels. Don't forget about case!

<!-- grade_id: isvowel, points: 0, grade: false, solution: true -->
```python
def isVowel(char):
    """Takes a character and returns True if it is a vowel, and False otherwise.
    
    Parameters
    ----------
    char : string
        string of length 1
        
    Returns
    -------
    boolean telling whether char is a vowel or not
    
    """
    # YOUR CODE HERE
    return any(char.lower()  == c for c in 'aeiou') 
```

<!-- grade_id: test_isvowel, points: 1, grade: true, solution: false -->
```python
"""(1 point) Test code for the previous function. This cell should NOT give any errors when it is run."""
assert isVowel('a')
assert isVowel('U')
assert not isVowel('m')
assert not isVowel(' ')
assert not isVowel('This is not a vowel')

print("Success!")
```

Output:
```
Success!
```

### 3. Reverse a string (1 point)

Define a function reverse() that computes the reversal of a string. For example, 
    >>> reverse("I am testing")
    "gnitset ma I"

<!-- grade_id: reverse, points: 0, grade: false, solution: true -->
```python
def reverse(string):
    """Reverse string and then return in.
    
    Parameters
    ----------
    string : str
        
        
    Returns
    -------
    string that is reversed
    """
    # YOUR CODE HERE
    return string[::-1]
```

<!-- grade_id: test_reverse, points: 1, grade: true, solution: false -->
```python
"""(1 point) Test code for the previous function. This cell should NOT give any errors when it is run."""
from nose.tools import assert_equal
assert_equal(reverse('I am testing'), 'gnitset ma I')
assert_equal(reverse('I am TesTinG'), 'GniTseT ma I')
assert_equal(reverse(''), '')

print("Success!")
```

Output:
```
Success!
```

### 4. Compute the length of a given list or string. (2 points)

Define a function that computes the length of a given list or string. (It is true that Python has the len() function built in, but do not use it for your solution — use a loop instead.)

For example:

    >>> length(['string1', 'string2'])
    2
    >>> length('string1')
    7

<!-- grade_id: length, points: 0, grade: false, solution: true -->
```python
def length(inputs):
    """Compute the length of a given list or string.
    
    Parameters
    ----------
    inputs : string or list
        
        
    Returns
    -------
    integer that is either the number of characters in the input string or the number of 
    entries in the input list.
    """
    # YOUR CODE HERE
    
    if inputs.count(inputs[-1]) > 1:
        return inputs.rindex(inputs[-1])+1
    else:
        return inputs.index(inputs[-1])+1
```

<!-- grade_id: test_length, points: 2, grade: true, solution: false -->
```python
"""(2 points) Test code for the previous function. This cell should NOT give any errors when it is run."""
from nose.tools import assert_equal
# Check a list
inputs = ['apple', 'orange']
assert_equal(length(inputs), 2)

# Check a string
inputs = 'tomato'
assert_equal(length(inputs), 6)

# What about a string with a space in it?
inputs = 'bacon and eggs'
assert_equal(length(inputs), 14)

# Also will look for not using the built-in len() function
len_old = len
len = None
inputs = 'tomato'
assert_equal(length(inputs), 6)
len = len_old

print("Success!")
```

Output:
```
Success!
```

### 5. Check for palindromes (2 points)

Define a function is_palindrome() that recognizes palindromes (i.e. words and phrases that look the same written forwards and backwards). Make sure that case and spaces are ignored in your function; for example,
    >>> is_palindrome("Never odd or even")
    True

You can use your solution to #3 for solving this problem.

<!-- grade_id: palindrome, points: 0, grade: false, solution: true -->
```python
def is_palindrome(phrase):
    """Returns True for a palindrome and False otherwise.
    
    Parameters
    ----------
    phrase : string
        
        
    Returns
    -------
    Boolean
    """
    # YOUR CODE HERE
    phrasenew=phrase.replace(" ","").lower()
    return phrasenew[::-1]==phrasenew
```

<!-- grade_id: test_palindrome, points: 2, grade: true, solution: false -->
```python
"""(2 points) Test code for the previous function. This cell should NOT give any errors when it is run."""
from nose.tools import assert_equal
assert is_palindrome("radar")
assert is_palindrome("rad dar")
assert is_palindrome("Never odd or even")
assert not is_palindrome("I am testing")

print("Success!")
```

Output:
```
Success!
```

### 6. Encode and decode using ROT-13 (2 points)

In cryptography, a Caesar cipher is a very simple encryption techniques in which each letter in the plain text is replaced by a letter some fixed number of positions down the alphabet. For example, with a shift of 3, A would be replaced by D, B would become E, and so on. The method is named after Julius Caesar, who used it to communicate with his generals. ROT-13 ("rotate by 13 places") is a widely used example of a Caesar cipher where the shift is 13. In Python, the key for ROT-13 may be represented by means of the following dictionary:

    key = {'a':'n', 'b':'o', 'c':'p', 'd':'q', 'e':'r', 'f':'s', 'g':'t', 'h':'u', 
           'i':'v', 'j':'w', 'k':'x', 'l':'y', 'm':'z', 'n':'a', 'o':'b', 'p':'c', 
           'q':'d', 'r':'e', 's':'f', 't':'g', 'u':'h', 'v':'i', 'w':'j', 'x':'k',
           'y':'l', 'z':'m', 'A':'N', 'B':'O', 'C':'P', 'D':'Q', 'E':'R', 'F':'S', 
           'G':'T', 'H':'U', 'I':'V', 'J':'W', 'K':'X', 'L':'Y', 'M':'Z', 'N':'A', 
           'O':'B', 'P':'C', 'Q':'D', 'R':'E', 'S':'F', 'T':'G', 'U':'H', 'V':'I', 
           'W':'J', 'X':'K', 'Y':'L', 'Z':'M'}

Your task in this exercise is to implement an encoder/decoder of ROT-13. Once you're done, you will be able to read the following secret message:

       Pnrfne pvcure? V zhpu cersre Pnrfne fnynq!

Note that since English has 26 characters, your ROT-13 program will be able to both encode and decode texts written in English. You should be able to do this using the properties of a dictionary.

<!-- grade_id: rot13, points: 0, grade: false, solution: true -->
```python
key = {'a':'n', 'b':'o', 'c':'p', 'd':'q', 'e':'r', 'f':'s', 'g':'t', 'h':'u', 
       'i':'v', 'j':'w', 'k':'x', 'l':'y', 'm':'z', 'n':'a', 'o':'b', 'p':'c', 
       'q':'d', 'r':'e', 's':'f', 't':'g', 'u':'h', 'v':'i', 'w':'j', 'x':'k',
       'y':'l', 'z':'m', 'A':'N', 'B':'O', 'C':'P', 'D':'Q', 'E':'R', 'F':'S', 
       'G':'T', 'H':'U', 'I':'V', 'J':'W', 'K':'X', 'L':'Y', 'M':'Z', 'N':'A', 
       'O':'B', 'P':'C', 'Q':'D', 'R':'E', 'S':'F', 'T':'G', 'U':'H', 'V':'I', 
       'W':'J', 'X':'K', 'Y':'L', 'Z':'M'}

def rot13(letter):
    """Swap letter indicated in key.
    
    Parameters
    ----------
    letter : string of length 1
        letter to be swapped using the key dictionary
        
    Returns
    -------
    letter
    """
    # YOUR CODE HERE
    keydecode=''
    for i in letter:
        if i.islower():
            keydecode+=key.get(i) 
        if i.isupper():
            i=i.lower()
            keydecode+=key.get(i).capitalize()
        if i not in key:
            keydecode+=i
        
    
    return keydecode
```

<!-- grade_id: test_rot13, points: 2, grade: true, solution: false -->
```python
"""(2 points) Test code for the previous function. This cell should NOT give any errors when it is run."""
phrase = 'Pnrfne pvcure? V zhpu cersre Pnrfne fnynq!'

translatedchars = list(map(rot13, phrase))  # Also used in #7, map connects a function with 
                                            # each letter in phrase.
translated = ''.join(translatedchars)  # we got a bunch of characters back. Join combines them
                                       # into a string

# recode
retranslatedchars = list(map(rot13, translated))
retranslated = ''.join(retranslatedchars)
    
# Check decode
from nose.tools import assert_equal
assert_equal(translated, "Caesar cipher? I much prefer Caesar salad!")
assert_equal(retranslated, phrase)  # check that it works both ways

print("Success!")
```

Output:
```
Success!
```

### 7. Fix the code (2 points)

"99 Bottles of Beer" is a traditional song in the United States and Canada. It is popular to sing on long trips, as it has a very repetitive format which is easy to memorize, and can take a long time to sing. The song's simple lyrics are as follows:

99 bottles of beer on the wall, 99 bottles of beer.
Take one down, pass it around, 98 bottles of beer on the wall.

The same verse is repeated, each time with one fewer bottle. The song is completed when the singer or singers reach zero.

Code to generate all verses of this song has already been written and is provided below. However, it has several problems in it that prevent it from running properly. Your task is to fix the code, keeping a list of what you change, so that the code generates 3 verses of the song.

You should simply remove the not implemented error and insert code part, and instead just fix the existing code for this problem.

<!-- grade_id: bottles, points: 0, grade: false, solution: true -->
```python
def words(num):
    """Returns words with proper number of beer bottles filled in."""
    
    if num != 1:
        return str(num) + ' bottles of beer on the wall, ' + str(num) + ' bottles of beer.\n' + \
            'Take one down, pass it around, ' + str(num-1) + ' bottles of beer on the wall.\n\n'

    elif num == 1:
        return str(num) + ' bottle of beer on the wall, ' + str(num) + ' bottle of beer.\n' + \
            'Take one down, pass it around, ' + str(num-1) + ' bottles of beer on the wall.'
        
def connect_words(numbottles):
    
    listofwords = []
    for i in range(numbottles, 0, -1):
        listofwords.append(words(i))
    
    return listofwords
    
# YOUR CODE HERE

```

<!-- grade_id: test_bottles, points: 2, grade: true, solution: false -->
```python
"""(2 points) Test code for the previous function. This cell should NOT give any errors when it is run."""

output = '''3 bottles of beer on the wall, 3 bottles of beer.
Take one down, pass it around, 2 bottles of beer on the wall.

2 bottles of beer on the wall, 2 bottles of beer.
Take one down, pass it around, 1 bottles of beer on the wall.

1 bottle of beer on the wall, 1 bottle of beer.
Take one down, pass it around, 0 bottles of beer on the wall.'''

numbottles = 3
from nose.tools import assert_equal
assert_equal(output, ''.join(connect_words(numbottles)))

print("Success!")
```

Output:
```
Success!
```

```python

```