```python
# Be sure to use Python 3
from nose.tools import assert_equal
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

# Core Language continued

### 1. Number lines in a text file

Write a function, `num_lines()`, that numbers all of the lines in a text file, skipping empty lines, and writes the output to another specified file. The output file might look something like

    1    This is the first line of the file
    2    This is the second line of the file
    
    3    This is the fourth line of the file

where the third line is empty, and not counted.

<!-- grade_id: lines, points: 0, grade: false, solution: true -->
```python
def num_lines(inputfilename, outputfilename):
    """Numbers all lines in a text file, skipping empty lines. Writes results to output file.
    
    Parameters
    ----------
    inputfilename : string
        Name of the input file to be read in.      
        
    Returns
    -------
    Nothing, but does write results to file outputfilename.
    """
    # YOUR CODE HERE
    f=open(inputfilename)
    ff=open(outputfilename,'w')
    f.seek(0) 
    n=0
    for line in f.readlines():                                        
        data = line.split()   
        if data==[]: 
            ff.write('\n')
        else: 
            ff.write(str(n)+' '+line)
            n+=1
     

```

```python
'\t\t\t\t    \n'.split()
```

Output:
```
[]
```

```python
f=open('inputfile.txt')
line= f.readlines()                                        
data = line[1].split(',')
data
```

Output:
```
['\n']
```

```python
num_lines('inputfile.txt', 'outputfile.txt')

```

<!-- grade_id: test_lines, points: 3, grade: true, solution: false -->
```python
"""(3 points) Test code for the previous function. This cell should NOT give any errors when it is run."""
# run code
inputfilename = 'inputfile.txt'  # this is in the hw directory
outputfilename = inputfilename.split('.')[0] + '_mod.txt'
num_lines(inputfilename, outputfilename)

# check for existence of output file
import os
assert os.path.exists(outputfilename)
from nose.tools import assert_equal

# check for text in output file being correct
fin = open(outputfilename, 'r')  # prepare to read
lines = fin.readlines()
line = lines[0].split(' ')
assert_equal(line[0], '0')  # check first word in line
assert_equal(line[-1], 'eros.\n')  # check final word in line
assert_equal(lines[1], '\n')  # check empty line
line = lines[2].split(' ')
assert_equal(line[0], '1')  # check first word in line
assert_equal(line[-1], 'sodales.\n')  # check final word in line
assert_equal(lines[3], '\n')  # check empty line
line = lines[4].split(' ')
assert_equal(line[0], '2')  # check first word in line
assert_equal(line[-1], 'a.\n')  # check final word in line
assert_equal(lines[5], '\n')  # check empty line
line = lines[6].split(' ')
assert_equal(line[0], '3')  # check first word in line
assert_equal(line[-1], 'ut.\n')  # check final word in line
assert_equal(lines[7], '\n')  # check empty line
line = lines[8].split(' ')
assert_equal(line[0], '4')  # check first word in line
assert_equal(line[-1], 'malesuada.\n')  # check final word in line
fin.close()
print('success!')
```

Output:
```
success!
```

### 2a-c. Convert list of words into list of integers of length of words

Write a series of functions that maps a list of words into a list of integers representing the lengths of the corresponding words. Write them in three different ways: 

1. using a for-loop

2. using the higher order function map(), and 

3. using list comprehensions. 

Each function will return a list of word lengths, for example,

    >>> method1(['list', 'of', 'words'])
    [4, 2, 5]

<!-- grade_id: forloop, points: 0, grade: false, solution: true -->
```python
# part a
def method1(input1):
    """Returns a list of integers representing the lengths of the corresponding input list of words.
    Do this using a for loop.
    
    Parameters
    ----------
    input1 : list of str
        inputs to be counted   
        
    Returns
    -------
    list of integers counting the elements in the input list
    """
    # YOUR CODE HERE
    nums=[]
    for wor in input1:
        nums.append(len(wor))
    return nums
```

<!-- grade_id: test_forloop, points: 1, grade: true, solution: false -->
```python
"""(1 point) Test code for the previous function. This cell should NOT give any errors when it is run."""
# check with a list
words = ['one', 'two', 'three']
assert_equal(method1(words), [3, 3, 5])
# Also will check that a for loop was used
print('success!')
```

Output:
```
success!
```

<!-- grade_id: map, points: 0, grade: false, solution: true -->
```python
# part b

def method2(input2):
    """Returns a list of integers representing the lengths of the corresponding input list of words.
    Do this using the function `map`.
    
    Parameters
    ----------
    input1 : list of str
        inputs to be counted   
        
    Returns
    -------
    list of integers counting the elements in the input list
    """
    # YOUR CODE HERE

    return list(map(len,input2))

```

<!-- grade_id: test_map, points: 1, grade: true, solution: false -->
```python
"""(1 point) Test code for the previous function. This cell should NOT give any errors when it is run."""
# check with a list
words = ['one', 'two', 'three']
assert_equal(method2(words), [3, 3, 5])
# Also will check that the map function was used
print('success!')
```

Output:
```
success!
```

<!-- grade_id: listcomprehension, points: 0, grade: false, solution: true -->
```python
# part c
def method3(input3):
    """Returns a list of integers representing the lengths of the corresponding input list of words.
    Do this using a list comprehension.
    
    Parameters
    ----------
    input1 : list of str
        inputs to be counted   
        
    Returns
    -------
    list of integers counting the elements in the input list
    """
    # YOUR CODE HERE
    nums= [len(word) for word in input3]
    return nums
```

<!-- grade_id: test_listcomprehension, points: 1, grade: true, solution: false -->
```python
"""(1 point) Test code for the previous function. This cell should NOT give any errors when it is run."""
# check with a list
words = ['one', 'two', 'three']
assert_equal(method3(words), [3, 3, 5])
# Also will check that a list comprehension was used
print('success!')
```

Output:
```
success!
```

<!-- grade_id: test_all, points: 1, grade: true, solution: false -->
```python
"""(1 point) Test code for the previous function. This cell should NOT give any errors when it is run."""
words = ['one', 'two', 'three']
# Check that the result is the same for all methods
assert_equal(method1(words), method2(words), method3(words))
print('success!')
```

Output:
```
success!
```

### 3. Reading in and parsing a file

Read in `GPS.dat`, a file containing GPS track information. This file is tab-delimited, but contains other characters as well. It has data from several drifters, whose names are listed in the file after the TRACK label; the first is named "ACTIVE LOG". 

a. Read in the date and time information from the file for the drifter ACTIVE LOG only, and form a list of datetime objects with this information, where each line in the file is a data point that should have its own datetime object. Your `get_datetimes()` function should return this list of datetime objects.

b. Separately, read in the locations and speed at all data points for each drifter. Create a dictionary with the name of each track as a key. The value of the dictionary will be another dictionary that has keys `latitude`, `longitude`, and `speed`.  That is, `gps['GANDALF']['speed']` will return a list containing a timeseries of the speed of that track.

<!-- grade_id: readandparse, points: 0, grade: false, solution: true -->
```python
# part a
from datetime import datetime

def get_datetimes(f):
    """Returns a list of datetime objects that have been read in from the input file object.
    
    Parameters
    ----------
    f : file object
        open file object to be read from  
        
    Returns
    -------
    list of datetime objects
    """
    # YOUR CODE HERE
    datetimes=[]
    f = open('GPS.dat')
    f.seek(0)  # This sets the pointer back to the beginning of the file. This allows us to run this
    for line in f.readlines():
        data = line.split('\t')   
        if not data: continue
        if data[0] == 'Trackpoint' and data[2] != '':    
            year = int(data[2][5:9])
            month = int(data[2][0:1])
            day = int(data[2][2:4])   
            hour = int(data[2][10])+12
            minute = int(data[2][12:14])
            second = int(data[2][15:17])
            dt = datetime(year, month, day, hour, minute, second)
            datetimes.append(dt)
#        if data[1] == 'LEGOLAS': break
            
    return datetimes
   

```

<!-- grade_id: test_readandparse, points: 4, grade: true, solution: false -->
```python
"""(4 points) Test code for the previous function. This cell should NOT give any errors when it is run."""
# tests for Part A
f = open('GPS.dat')
f.seek(0)  # This sets the pointer back to the beginning of the file. This allows us to run this
           # block of code many times without reopening the file each time.
datetimes = get_datetimes(f)
# Check results
assert datetimes[0] == datetime(2006, 5, 20, 13, 34, 55)  # check first entry
assert datetimes[100] == datetime(2006, 5, 20, 13, 59, 57)
assert datetimes[200] == datetime(2006, 5, 20, 14, 24, 58)
assert datetimes[-50] == datetime(2006, 5, 20, 15, 19, 46)
assert datetimes[-1] == datetime(2006, 5, 20, 15, 31, 48)  # check last entry
assert len(datetimes) == 469
print('success!')
```

Output:
```
success!
```

<!-- grade_id: readandparse_grad, points: 0, grade: false, solution: true -->
```python
# part b

missing_value = -999

def get_gps(f):
    """Read in the locations and speed at all data points for each drifter in file object, f.
    
    Parameters
    ----------
    f : file object
        open file object to be read from  
        
    Returns
    -------
    A dictionary of a dictionary for each drifter, which each contains keys `latitude`, `longitude`, and `speed` 
    """
    # YOUR CODE HERE

    f.seek(0)  # This sets the pointer back to the beginning of the file. This allows us to run this
    datetimes=[]
    tracks = {}
    for line in f.readlines():
        data = line.split()
    #if data[1] = 'ACTIVE LOG': continue        
        if not data: continue
        if data[0] == 'Track':
            trackname = data[1]
            tracks[trackname] = {}
            tracks[trackname]['lon'] = []
            tracks[trackname]['lat'] = []
            tracks[trackname]['speed'] = []
        if data[0] == 'Trackpoint': 
            lat = int(data[1][1:]) + float(data[2])/60   # the index to data[1] is used to trim the 'N' off of the string.
            tracks[trackname]['lat'].append(lat)
            lon = int(data[3][1:]) + float(data[4])/60   
            tracks[trackname]['lon'].append(lon)
            timeorspeed = data[6]
#         if trackname == 'ACTIVE':
            if  len(data)>=13:
                speed = float(data[13])
            else:
                speed = -999    
            tracks[trackname]['speed'].append(speed)
    tracks['ACTIVE LOG'] = tracks['ACTIVE'] 
    del tracks['ACTIVE']
    return tracks
```

```python
d1 = {}
d1['ACTIVE LOG'] = {}
d1['ACTIVE LOG']['lon'] = []
d1['ACTIVE LOG']['lat'] = []

d1['ACTIVE LOG']['lon'].append(42.5)
d1['ACTIVE LOG']['lat'].append(42.5)
d1['ACTIVE'] = d1['ACTIVE LOG'] 
d1
```

Output:
```
{'ACTIVE': {'lat': [42.5], 'lon': [42.5]},
 'ACTIVE LOG': {'lat': [42.5], 'lon': [42.5]}}
```

<!-- grade_id: test_readandparse_grad, points: 5, grade: true, solution: false -->
```python
"""(5 points) Test code for the previous function. This cell should NOT give any errors when it is run."""
# tests for Part B
from nose.tools import assert_equal
f = open('GPS.dat')
f.seek(0)  # This sets the pointer back to the beginning of the file. This allows us to run this
           # block of code many times without reopening the file each time.

gps = get_gps(f)
    
assert_equal(len(gps), 5)  # test for all drifters being there
assert gps['STRIDER']['lon']  # test for one containing lon
assert_equal(len(gps['SAMWISE']), 3)  # test for 3 entries in SMA
assert_equal(gps['LEGOLAS']['lat'][5], 42 + 48.872/60)
assert_equal(gps['ACTIVE LOG']['speed'][-5], 4.7)

import numpy as np  # using numpy for this test
assert np.allclose(gps['SAMWISE']['lat'], [42.82065, 42.819833333333335, 42.81883333333333])
assert np.allclose(gps['SAMWISE']['lon'], [70.76708333333333, 70.76886666666667, 70.77156666666667])
assert np.allclose(gps['SAMWISE']['speed'], [-999, -999, -999])
print('success!')
```

Output:
```
success!
```

### 4. Constructing terms in a Taylor series to approximate a function

The Taylor series expansion for the sine function is given by 
$$ \text{sin}(x) = \Sigma^\infty_{k=0} \frac{(-1)^k x^{1+2k}}{(1+2k)!}$$

This expression gives an exact form for sin(x) when an infinite number of terms are used; any truncation gives an approximation. Write a loop that constructs terms from $k=0$ up to $k=N$. Don't use any function for the factorial since we haven't gotten to NumPy... create it yourself, instead. Provide the value of sin(0.5) with N=5 terms. How many terms do you need to include before the answer starts to plateau?


<!-- grade_id: taylor, points: 0, grade: false, solution: true -->
```python
import math
def calc_sine(x, N):
    """Approximate sin(x) using Taylor series with N terms.
    
    Parameters
    ----------
    x : float
        value to approximate sine function at
    N : integer
        number of terms +1 to approximate, from k=0 to k=N
        
    Returns
    -------
    approximation of sine(x) to N+1 terms
    """
    # YOUR CODE HERE
    tot=0
    for k in range(N):
        t=(((-1)**k)*(x**(1+(2*k))))/(math.factorial(2*k+1))
        tot+=t
    return tot
```

<!-- grade_id: test_taylor, points: 3, grade: true, solution: false -->
```python
"""(3 points) Test code for the previous function. This cell should NOT give any errors when it is run."""
x = 0.5
N = 5
tot = calc_sine(x, N)
import numpy as np  # using numpy for this test
assert np.allclose(tot, 0.4794255386041834)
print('success!')
```

Output:
```
success!
```

### 5. Create a 'line' class

Create a class, 'line' that takes a slope and intercept as inputs to create the class instance, where these value are stored. The class should have one method, 'evaluate' that calculates the y-value of the line given x. For example:

    >>> line(0.5, 2.0)  # slope of 1/2, y-intercept of 2
    >>> line.evaluate(3.0)
    3.5

<!-- grade_id: line, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
class line(object):
    def __init__(self, x, y):
        self.x=x
        self.y=y
        
    def evaluate(self,input2):
        return self.x*input2+self.y

p1=line(0.5,2.0)
p1.evaluate(3.0)
```

Output:
```
3.5
```

<!-- grade_id: test_line, points: 3, grade: true, solution: false -->
```python
"""(3 points) Test code for the previous function. This cell should NOT give any errors when it is run."""
import numpy as np  # using numpy for this test
assert np.allclose(line(0.5, 2.0).evaluate(3.0), 3.5)
assert np.allclose(line(-0.5, 2.0).evaluate(3.0), 0.5)
assert np.allclose(line(0., 3.0).evaluate(3.0), 3.0)
print('success!')
```

Output:
```
success!
```

```python

```